<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 476168  -->
  <question type="stack">
    <name>
      <text>000 Start 0.2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<script src="https://marvin.hs-bochum.de/~mneugebauer/alquiz-qpool-instant-tutoring.js"></script>
<script>
    ALQuiz.setCurrentQuestionId("start");
</script>
<p class="hint">Hey! Ich helfe dir beim Lösen der Mathe-Aufgaben. In diesem Übungsraum leite ich dich durch 5 Mathe-Welten. Jede Mathe-Welt besteht aus mehreren Hilfsschritt-Leveln und einem Endgegner <img src="https://marvin.hs-bochum.de/~mneugebauer/skull.svg" style="width:1em">. Das Ziel ist, in jeder Welt die Endgegner-Aufgabe zu lösen. Du brauchst die Hilfsschritt-Level nicht unbedingt lösen. Doch wenn du sie löst, wird es dir leichter fallen, die Endgegner-Aufgabe zu lösen.<br><br>Nutze die <span class="show-on-mobile-only"><a href="javascript:;" onclick="let showsidebarbutton = document.getElementById('showsidebaricon'); if (showsidebarbutton != undefined) { showsidebarbutton.click(); return; } let questionCard = document.querySelector('[id*=quiznavbutton]'); let drawerParent = questionCard.closest('.drawer'); if (drawerParent != undefined) { let target = drawerParent.id; let drawerOpener = document.querySelector('button[data-target='+target+'][data-action=toggle]'); if (drawerOpener != undefined) { drawerOpener.click(); } }">Navigation</a></span><span class="show-on-desktop-only">Navigation (rechts)</span> um dich durch die Level und Welten zu bewegen.<br><br>Mit Klick auf <a href="javascript:;" onclick="tutorialFocusElement(document.querySelector('.submitbtns'));">"Nächste Frage"</a>
    <!--liefere ich dir automatisch die nächste Frage, die zu deinem Lernstand passt.-->Geht es los.<br>Viel Erfolg!<br>
    Hier kannst du schon einmal die Eingabe ausprobieren. Versuche doch einfach Mal, die Zahl {@tasklist[1]@} einzugeben und deine Eingabe mit "Enter" überprüfen zu lassen.<br>
</p>

<p></p>
<p>Machen Sie sich mithilfe des untenstehenden Eingabefelds mit der mathematischen Eingabe per Tastatur vertraut. Probieren Sie zum Beispiel <code>42</code>, <code>2*x+1</code>, <code>3/5</code> oder <code>sqrt(2)</code>. Schreiben Sie einen Ausdruck pro Zeile.</p>
Durch das Beginnen einer neuen Zeile (Enter) wird Ihre Eingabe vom Digitalen Mentor geprüft.<br>
<p></p>
<p>[[input:ans1]][[validation:ans1]][[feedback:prt1]]</p><br>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>tasklist:[42,2*x+1,3/5,sqrt(2)];
i:1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>equiv</type>
      <tans>tasklist</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>hideequiv</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[1]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Genau! So kannst du eine Eingabe machen und sie prüfen lassen. Gibt jetzt darunter {@tasklist[2]@} ein.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[2]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Genau! So kannst du eine Eingabe machen und sie prüfen lassen. Gibt jetzt {@tasklist[3]@} ein.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[3]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Genau! So kannst du eine Eingabe machen und sie prüfen lassen. Gibt jetzt {@tasklist[4]@} ein.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[4]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Prima, das soll für den Einstieg reichen! Viel Erfolg! Mit Klick auf "Nächste Frage" geht es los.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.1</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Genau, so kannst du eine Eingabe machen und sie prüfen lassen. Versuche doch noch einmal, einen der folgenden Ausdrücke einzugeben.</p><p dir="ltr" style="text-align: left;">{@tasklist[1]@}</p><p dir="ltr" style="text-align: left;">{@tasklist[2]@}</p><p dir="ltr" style="text-align: left;">{@tasklist[3]@}</p><p dir="ltr" style="text-align: left;">{@tasklist[4]@}<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>
