<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 476168  -->
  <question type="stack">
    <name>
      <text>000 Start 0.2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<script src="https://marvin.hs-bochum.de/~mneugebauer/alquiz-qpool-instant-tutoring-en.js"></script>
<script>
    ALQuiz.setCurrentQuestionId("start");
</script>
<p class="hint">Hey! I'll help you solve math problems. In this training area I will guide you through 5 math worlds. Each math world consists of several help step levels and a boss <img src="https://marvin.hs-bochum.de/~mneugebauer/skull.svg" style="width:1em">. The aim is to solve the boss task in each world. You do not necessarily need to solve the help step levels. But if you solve them, it will be easier for you to solve the boss task.<br><br>Use the <span class="show-on-mobile-only"><a href="javascript:;" onclick="let showsidebarbutton = document.getElementById('showsidebaricon'); if (showsidebarbutton != undefined) { showsidebarbutton.click(); return; } let questionCard = document.querySelector('[id*=quiznavbutton]'); let drawerParent = questionCard.closest('.drawer'); if (drawerParent != undefined) { let target = drawerParent.id; let drawerOpener = document.querySelector('button[data-target='+target+'][data-action=toggle]'); if (drawerOpener != undefined) { drawerOpener.click(); } }">navigation</a></span><span class="show-on-desktop-only">navigation (rechts)</span> to move through the levels and worlds.<br><br>Click on <a href="javascript:;" onclick="tutorialFocusElement(document.querySelector('.submitbtns'));">"Next Question"</a>
    <!--liefere ich dir automatisch die nächste Frage, die zu deinem Lernstand passt.--> to get started.<br>Good luck!<br>
    You can try out the input here. Just try entering the number {@tasklist[1]@} and pressing "Enter" to check your input.<br>
</p>

<p></p>
<p>Use the input field below to familiarize yourself with mathematical input using the keyboard. For example, try <code>42</code>, <code>2*x+1</code>, <code>3/5</code> or <code>sqrt(2)</code>. Write one expression per line.</p>
By starting a new line (Enter), your input will be checked by the Digital Mentor.<br>
<p></p>
<p>[[input:ans1]][[validation:ans1]][[feedback:prt1]]</p><br>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>tasklist:[42,2*x+1,3/5,sqrt(2)];
i:1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct. Well done!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>equiv</type>
      <tans>tasklist</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>hideequiv</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[1]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Exactly! So you can make an input and have it checked. Now enter {@tasklist[2]@} underneath.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[2]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Exactly! So you can make an input and have it checked. Now enter {@tasklist[3]@} underneath.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[3]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Exactly! So you can make an input and have it checked. Now enter {@tasklist[4]@} underneath.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>last(ans1)</sans>
        <tans>tasklist[4]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Great, that should be enough to get you started! Good luck! Click on "Next question" to get started.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.1</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Exactly, so you can make an input and have it checked. Try entering one of the following expressions again.</p><p dir="ltr" style="text-align: left;">{@tasklist[1]@}</p><p dir="ltr" style="text-align: left;">{@tasklist[2]@}</p><p dir="ltr" style="text-align: left;">{@tasklist[3]@}</p><p dir="ltr" style="text-align: left;">{@tasklist[4]@}<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>
